Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Week2 Assignments


This assignment has three components. Each of the first three components receives one point if all the tests pass. 
The last component on Pandas may require more effort, it consists of three subtasks, each of the subtasks receives 1 point.

All the exercises are designed so that the solutions will need only one or a few lines of code.

Some concepts may be new to you and may require digging into the Python, NumPy or Pandas documentation, the links are provided.

Do not hesitate to contact instuctors and TA via #week2 channel on Slack if you get stuck. Join the channel first by clicking on Channels.

## Part A. Create a missing function (1 point)

In this exercise you need to create a function __missing_link(x)__ that is passed to another function as an argument in order to perform a calculation.

We know the final result (see the assert operator), but we do not know the intermediate calculation leading to that result.

Read about Python built-in functions __all()__ and __zip()__
https://docs.python.org/3.3/library/functions.html

and about the iterators and generators here:
https://docs.python.org/3.3/library/stdtypes.html#typeiter

```python
def calculate(func, it):
    """
    Performs calculation by applying *func* to each item of iterator *it*
    Returns a generator as a result.
    """
    return (2 * func(a) for a in it)

```

```python
list(calculate(lambda x: x**2, range(7)))
```

Output:
```
[0, 2, 8, 18, 32, 50, 72]
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
def missing_link(x):
    """Define a function that will be passed to calculate() as an argument"""
    return x**2
## You can check the result of the missing_link() function and of calculate() if you wish:
# print(list(map(missing_link, range(5))))
# print(list(calculate(missing_link, range(5))))
```

```python
missing_link(2.5)
```

Output:
```
6.25
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
_observed_results = calculate(missing_link, range(7))
_expected_results = [0, 2, 8, 18, 32, 50, 72]

assert all(a == b for a, b in zip(_observed_results, _expected_results))
```

## Part B. NumPy (1 point)

Define __x__ as a subdivision of an interval from -4 PI to 4 PI into 32 equal parts, i.e. with a PI/4 step.

Including both endpoints that should give 33 points.

Using NumPy calculate __cos()__ and __sin()__ and find the values of

__x__ where __cos(x)__ is equal to __sin(x)__ and store these values in the variable __y__.

Use NumPy vector operations.

Use __np.pi__ constant and __np.linspace()__ function: 
https://docs.scipy.org/doc/numpy/reference/generated/numpy.linspace.html

Note that due to the way floating points are stored in memory exact comparison is nearly always impossible. You should use __np.isclose()__ instead. That would allow some room for floating point errors.
https://docs.scipy.org/doc/numpy-dev/reference/generated/numpy.isclose.html

This plot may be helpful:
http://www.wolframalpha.com/input/?i=plot+sinx+and+cosx+from+-4pi+to+4pi

```python
!conda install numpy
```

```python
!conda list | grep numpy
```

Output:
```
numpy                     1.16.4           py37h926163e_0  
numpy-base                1.16.4           py37ha711998_0
```

<!-- grade_id: cell-8d46731c6bc218b6, points: 0, grade: false, solution: true -->
```python
import numpy as np
```

<!-- grade_id: cell-8d46731c6bc218b6, points: 0, grade: false, solution: true -->
```python
x = np.linspace(-4*np.pi, 4*np.pi, num = 33)
y = x[np.isclose(np.cos(x), np.sin(x))]

# YOUR CODE HERE
pass
```

```python
len(x)
```

Output:
```
33
```

```python
np.cos(x)
```

Output:
```
array([ 1.00000000e+00,  7.07106781e-01, -4.28626380e-16, -7.07106781e-01,
       -1.00000000e+00, -7.07106781e-01,  3.06161700e-16,  7.07106781e-01,
        1.00000000e+00,  7.07106781e-01, -1.83697020e-16, -7.07106781e-01,
       -1.00000000e+00, -7.07106781e-01,  6.12323400e-17,  7.07106781e-01,
        1.00000000e+00,  7.07106781e-01,  6.12323400e-17, -7.07106781e-01,
       -1.00000000e+00, -7.07106781e-01, -1.96005386e-15,  7.07106781e-01,
        1.00000000e+00,  7.07106781e-01, -1.47019514e-15, -7.07106781e-01,
       -1.00000000e+00, -7.07106781e-01, -2.20498322e-15,  7.07106781e-01,
        1.00000000e+00])
```

```python
x.shape
```

Output:
```
(33,)
```

<!-- grade_id: cell-c7edebc7b2cd1c7b, points: 1, grade: true, solution: false -->
```python
assert x.shape[0] == 33
assert -4*np.pi in x
assert 0.0 in x
assert 4*np.pi in x

assert y.shape[0] == 8
assert np.all(np.isclose(y/np.pi, np.array([-3.75, -2.75, -1.75, -0.75,  0.25,  1.25,  2.25,  3.25])))
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-28-6ec04205bb22> in <module>
      4 assert 4*np.pi in x
      5 
----> 6 assert y.shape[0] == 8
      7 assert np.all(np.isclose(y/np.pi, np.array([-3.75, -2.75, -1.75, -0.75,  0.25,  1.25,  2.25,  3.25])))

AssertionError:
```

## Part C. Working with Pandas dataframes (3 points)

We will explore FBI reports on gun checks provided by the National Instant Criminal Background Check System (NICS)
https://www.fbi.gov/services/cjis/nics

Before ringing up the sale, cashiers call in a check to the FBI or to other designated agencies to ensure that each customer does not have a criminal record or isn’t otherwise ineligible to make a purchase. More than 230 million such checks have been made, leading to more than 1.3 million denials.

NICS and background checks is a hot topic and it is important to be able to do some basic fact-checking using the data available. https://www.washingtonpost.com/news/fact-checker/wp/2020/02/23/fact-checking-trump-nra-claims-on-gun-background-checks/?utm_term=.3e0284ad3774

The FBI NICS provides data as PDF reports, which is a really bad example of distributing the data.
There is a community-developed parser that extracted the data from PDF files. Parsed dataset that we will be using is available here: 
https://github.com/BuzzFeedNews/nics-firearm-background-checks/blob/master/README.md

Note that the number of background checks can not be directly interpreted as the number of guns sold because the actual sale protocols vary state to state.

For reproducibility we will be using a local copy of the file.

```python
import pandas as pd

# NICS parsed dataset url (local copy)
url = "https://github.com/biof509/spring2019/blob/master/week2/nics-firearm-background-checks.csv?raw=true"
guns = pd.read_csv(url)
```

```python
len(guns.columns.to_list())
```

Output:
```
27
```

```python
guns.shape
```

Output:
```
(13090, 27)
```

```python
guns.describe()
```

Output:
```
permit  permit_recheck        handgun       long_gun  \
count   13066.000000     1705.000000   13070.000000   13071.000000   
mean     6852.824966     1756.825806    6153.306809    7811.399434   
std     25796.442361    14385.751256    8900.603197    9250.726906   
min         0.000000        0.000000       0.000000       0.000000   
25%         0.000000        0.000000     925.000000    2102.500000   
50%       602.000000        0.000000    3174.500000    5130.000000   
75%      4633.500000        2.000000    7543.750000   10400.500000   
max    522188.000000   199766.000000  107224.000000  108058.000000   

              other      multiple         admin  prepawn_handgun  \
count   6105.000000  13090.000000  13067.000000     11147.000000   
mean     399.237510    273.582735     57.232724         4.913250   
std     1325.676832    775.608156    592.143131        11.002679   
min        0.000000      0.000000      0.000000         0.000000   
25%       20.000000     15.000000      0.00
```

```python
guns.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 13090 entries, 0 to 13089
Data columns (total 27 columns):
month                        13090 non-null object
state                        13090 non-null object
permit                       13066 non-null float64
permit_recheck               1705 non-null float64
handgun                      13070 non-null float64
long_gun                     13071 non-null float64
other                        6105 non-null float64
multiple                     13090 non-null int64
admin                        13067 non-null float64
prepawn_handgun              11147 non-null float64
prepawn_long_gun             11145 non-null float64
prepawn_other                5720 non-null float64
redemption_handgun           11150 non-null float64
redemption_long_gun          11149 non-null float64
redemption_other             5720 non-null float64
returned_handgun             2805 non-null float64
returned_long_gun            2750 non-null float64
returned_other   
```

```python
guns.loc[0:4]
```

Output:
```
month       state   permit  permit_recheck  handgun  long_gun   other  \
0  2020-08     Alabama  22465.0             0.0   5991.0    5741.0   218.0   
1  2020-08      Alaska    218.0             0.0   2593.0    3057.0   220.0   
2  2020-08     Arizona  12014.0           455.0  10833.0    6987.0   887.0   
3  2020-08    Arkansas   5684.0           690.0   4809.0    5590.0   226.0   
4  2020-08  California  38687.0             0.0  34617.0   25666.0  3117.0   

   multiple  admin  prepawn_handgun  ...  returned_other  rentals_handgun  \
0       274    0.0             12.0  ...             0.0              0.0   
1       181    0.0              5.0  ...             0.0              0.0   
2       693    0.0             18.0  ...             1.0              0.0   
3       342   95.0             14.0  ...             0.0              0.0   
4         0    0.0              0.0  ...             0.0              0.0   

   rentals_long_gun  private_sale_handgun  private_sale_long_gun  \
0    
```

```python
guns.head()
```

Output:
```
month       state   permit  permit_recheck  handgun  long_gun   other  \
0  2020-08     Alabama  22465.0             0.0   5991.0    5741.0   218.0   
1  2020-08      Alaska    218.0             0.0   2593.0    3057.0   220.0   
2  2020-08     Arizona  12014.0           455.0  10833.0    6987.0   887.0   
3  2020-08    Arkansas   5684.0           690.0   4809.0    5590.0   226.0   
4  2020-08  California  38687.0             0.0  34617.0   25666.0  3117.0   

   multiple  admin  prepawn_handgun  ...  returned_other  rentals_handgun  \
0       274    0.0             12.0  ...             0.0              0.0   
1       181    0.0              5.0  ...             0.0              0.0   
2       693    0.0             18.0  ...             1.0              0.0   
3       342   95.0             14.0  ...             0.0              0.0   
4         0    0.0              0.0  ...             0.0              0.0   

   rentals_long_gun  private_sale_handgun  private_sale_long_gun  \
0    
```

```python
# Use .head() .info() and .describe() to explore the dataset
```

### Part C. Subtask 1 (1 point)

First, use __pd.to_datetime()__ with argument __yearfirst=True__ to convert the column __"month"__ to a Pandas Series with DateTime objects. Add a new column __"year"__ to __guns__ dataframe and save the results of conversion there.
https://pandas.pydata.org/pandas-docs/stable/generated/pandas.to_datetime.html

You can access Python __datetime.date__ objects via the __.dt__ property of Pandas Series:
https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.dt.date.html

Look up the attributes of __datetime.date__ class, we will need attribute __.year__
https://docs.python.org/3/library/datetime.html


```python
df.loc[]
```

```python
guns['year'] = pd.to_datetime(guns['month'], yearfirst=True).dt.year
```

<!-- grade_id: cell-377e7ae405383f99, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass

```

<!-- grade_id: cell-fd44161086b39a45, points: 1, grade: true, solution: false -->
```python
assert (guns['year'].min(), guns['year'].max()) == (1998, 2020)
```

### Part C. Subtask 2 (1 point)

Group __guns__ dataframe by year and sum up the __totals__ (together, regardless of state). Use the variables
__totals_2000__ and __totals_2017__ to store the corresponding results.

You will need https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.groupby.html

<!-- grade_id: cell-75001ba281dc95cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```

<!-- grade_id: cell-b7846271cdcb45c0, points: 1, grade: true, solution: false -->
```python
assert totals_2000 == 8427096
assert totals_2017 == 24955919
```

### Part C. Subtask 3 (1 point)

Group data by state (regardless of year) and calculate the mean value of __long_gun__ and __handgun__ checks separately for each state. Calculate the number of states that had more long gun background checks on average over the years than handgun checks. Calculate the number of states with more handgun checks. Store these results in __states_with_more_long_guns__ and __states_with_more_handguns__ variables, respectively.

Hint: Use vector operations. No for loops are needed. A result of comparison of two vectors is a vector of booleans. You can sum up the vector of booleans to calculate the number of True values in it.

<!-- grade_id: cell-48eea5c5fac136c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pass
```

<!-- grade_id: cell-ec454c9539fb239e, points: 1, grade: true, solution: false -->
```python
assert states_with_more_long_guns == 35
assert states_with_more_handguns == 20
```

```python

```