# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

```python
"""This is a docstring!"""
```

Output:
```
'This is a docstring!'
```

$\beta=5$

```python
alpha = 4
```

```python
%who
```

Output:
```
SVC	 X	 alpha	 classifier	 datasets	 df	 iris	 pd	 predictions	 
train_test_split	 var_list	 y
```

```python
var_list = %who_ls
```

```python
var_list
```

Output:
```
['SVC',
 'X',
 'alpha',
 'classifier',
 'datasets',
 'df',
 'iris',
 'pd',
 'predictions',
 'train_test_split',
 'var_list',
 'y']
```

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
import pandas as pd
```

```python
df = pd.DataFrame(iris.data, columns=iris.feature_names)
```

look at the first couple of rows with head 

```python
df.head()
```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)
0                5.1               3.5                1.4               0.2
1                4.9               3.0                1.4               0.2
2                4.7               3.2                1.3               0.2
3                4.6               3.1                1.5               0.2
4                5.0               3.6                1.4               0.2
```

index the column names

```python
df.columns[0]
```

Output:
```
'sepal length (cm)'
```

```python
df['sepal length (cm)'].head()
```

Output:
```
0    5.1
1    4.9
2    4.7
3    4.6
4    5.0
Name: sepal length (cm), dtype: float64
```

```python
df.columns = [s.replace(" ", "_") for s in df.columns]
```

```python
df.columns = [s.replace("(", "").replace(")", "") for s in df.columns]
```

```python
df.columns
```

Output:
```
Index(['sepal_length_cm', 'sepal_width_cm', 'petal_length_cm',
       'petal_width_cm'],
      dtype='object')
```

```python
iris.feature_names
```

Output:
```
['sepal length (cm)',
 'sepal width (cm)',
 'petal length (cm)',
 'petal width (cm)']
```

iris. "tab" will show everything iris contains

magic commands can be found using ls magic ?maybe?

```python
%%bash 
echo hello world
```

Output:
```
hello world
```

who_ls will show all of the variables you have created 

```python
%who_ls
```

Output:
```
['SVC',
 'X',
 'alpha',
 'classifier',
 'datasets',
 'df',
 'iris',
 'pd',
 'predictions',
 'train_test_split',
 'var_list',
 'y']
```

```python
iris.target_names
```

Output:
```
array(['setosa', 'versicolor', 'virginica'], dtype='<U10')
```

```python
iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-96-eb43bd51aa77> in <module>
      3 # Load the iris dataset
      4 # YOUR CODE HERE
----> 5 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
%who
```

Output:
```
SVC	 X	 alpha	 classifier	 datasets	 df	 iris	 pd	 predictions	 
train_test_split	 var_list	 y
```

```python
# support vector classifier is SVC --> import 
from sklearn.svm import SVC
```

classification - outcome = discrete
regression - outcome = continuous

```python
X = iris.data
y = iris.target
```

```python
# instantiate 
classifier = SVC()
```

```python
# skl requires two dimensions
X.shape
```

Output:
```
(150, 4)
```

```python
# fit
classifier.fit(X, y)
```

```python
#predict
predictions = classifier.predict(X)
```

Output:
```
/Users/yaklichel/anaconda3/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
```

The below will give a classification accuracy 

```python
classifier.score(X,y)# shift tab will give the usage
```

Output:
```
0.9866666666666667
```

```python
svc = SVC()
```

```python
svc
```

Output:
```
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto_deprecated',
    kernel='rbf', max_iter=-1, probability=False, random_state=None,
    shrinking=True, tol=0.001, verbose=False)
```

CamelCase
snake_case

- Overfitting (Variance): tailored to a specific dataset (does not generalize well), misses the point, memorizes the data
- Underfitting (Bias): accuracy could be better, too general 

1. Import
2. Instantiate (a class)
  - Estimators:
    - Classifiers
    - Regressors
3. Fit
4. Predict

```python
predictions
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
from sklearn.model_selection import train_test_split
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```